DOCUMENT();

loadMacros(
    'PGstandard.pl',             'PGML.pl',
    'parserRadioMultiAnswer.pl', 'PGcourse.pl'
);

$a = random(2, 4);
$b = random(3, 7);

$rma = RadioMultiAnswer(
    [
        [
            'The solution set is \(\big\{\) %s \(\big\}\).  '
                . 'The equation is a conditional equation.',
            List(0)
        ],
        [
            'The solution set is \(\big\{x|x\neq\) %s \(\big\}\).  '
                . 'The equation is an identity.',
            List(0)
        ],
        [
            'The solution set is the set of real numbers.  The equation is an identity.'
        ],
        [
            'The solution set is the empty set.  The equation is inconsistent.'
        ]
    ],
    2,
    checkTypes => 0
)->with(
    checker => sub {
        my ($correct, $student, $self, $ans) = @_;
        if ($correct->[0] == $student->[0]) {
            $rma->appendMessage('Good job!');
            return 1;
        } else {
            $rma->appendMessage('Wrong!');
            return 0;
        }
    }
);

BEGIN_PGML
Solve the following equation.  Then identify the equation as an identity, an inconsistent
equation, or a conditional equation.

>>[`[$a](x - [$b]) = [$a] x - [$a * $b]`]<<

Select the correct choice below, and fill in the answer box if necessary.

[_____]{$rma}
END_PGML

ENDDOCUMENT();
